<a href="https://colab.research.google.com/github/Liuyang1018/introml/blob/main/exercise-11.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

<div class="alert-block alert-info alert"> Hello! If this is your first assignment please read the following instructions:

Read the questions carefully and code what is asked in the respective code block whenever you see a ``` #TODO: YOUR CODE HERE. ```
It is important to do the Coding tasks IN ORDER, as the majority of coding blocks are dependent on the
previous coding blocks in the notebook.
Please also note that there is a **raise NotImplementedError()** for each task or function you have to
implement, so don't forget to remove it before testing your code, otherwise you will always get a "not
implemented" error.

You need to manually upload the data used in this exercise to Colab. Just create a folder 'data' and paste the respective files into this folder. Otherwise, you will get a 'File not found' error. You can find the data in StudOn.

Aside from coding blocks, there are also testing blocks to check whether your implementations
are correct, you can run them but do not change the code inside of them.
The tests are determined via the ASSERT STATEMENTS that you see on the testing blocks, which can help you
out during the coding of the exercises. If you pass all the visible asserts, you will probably get a good
grade on the exercise, just remember that there are also hidden tests, so hardcoding the assert statements
won't help you.

REMEMBER: This Colab notebook DOES NOT HAVE AN AUTOSAVE FUNCTION, so whenever you want to close it, please
remember to save it by clicking on the "file" tab in the upper left corner of the page, followed by
clicking on the "Save a copy to GitHub" button and selecting your corresponding GitHub repository:
"madlab-biosig/exercise-x-YourGitHubUsername"

Before you save the final version of your assignment, please make sure everything runs as expected.
First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all
cells** (in the menubar, select Cell$\rightarrow$Run All).
</div>

### BioSig Exercises WS 23/24
---

# Exercise 11

```python
try:
    %matplotlib widget
    print('Using widgets')
except:
    try:
        print('Using Notebook widgets')
        %matplotlib notebook
    except:
        print('Using static backend')
        %matplotlib inline

import matplotlib.pyplot as plt
plt.rcParams.update({'figure.max_open_warning': 40})
import pandas as pd
import numpy as np
import scipy
import scipy.signal
```

Output:
```
Using Notebook widgets
Warning: Cannot change to a different GUI toolkit: notebook. Using widget instead.
Using static backend
```

## Ex 11.1: Extraction of Q-, R-, and S-peaks

To perform waveform analysis on individual heartbeats in an ECG signal, we need to find all relevant peaks within each repetition of the signal.

a) Load the example ecg file (`103.csv`) from the data folder using pandas. The signal is sampled at 360 Hz. Optionally apply a band-pass filter to see its effect.

b) Load the file `103_all.csv`. It contains the position of all R-peaks in the signal.

c) Extract all QRST complexes from the signal by considering a region spanning from 50 ms before until 250 ms after the R-peak.

d) Create a function to find the Q-peak by looking for a minimum left of the R-peaks within a reasonable interval.

e) Create a function to find the S-peak by looking for a minimum right of the R-peaks within a reasonable interval.

f) Add the position of the Q- and the S-peak as new columns to the DataFrame loaded in step b)


<!-- grade_id: cell-1729cc2cb7f30c9a, points: 0, grade: false, solution: true -->
```python
# Load the data and store it in a pd.Series called signal.
# Ensure that the index is in seconds
SUBJECT = 103
sampling_rate = 360  # hz

# YOUR CODE HERE
file_path = f'data/{SUBJECT}.csv'
signal = pd.read_csv(file_path, header=0,names=['index','ecg'],index_col='index')


signal.index = np.arange(0, len(signal)) / sampling_rate
# Optionally apply a bandpass filter to the signal (3-50 Hz passband)
# Test all computations in this exercise  with and without the filter
# Store the filter coefficients as 'b' and 'a'
# YOUR CODE HERE
lowcut = 3.0  # Lower cutoff frequency in Hz
highcut = 50.0  # Upper cutoff frequency in Hz
b, a = scipy.signal.butter(N=5, Wn=[lowcut, highcut], fs=sampling_rate,btype='band')

signal['ecg'] = scipy.signal.filtfilt(b, a,signal['ecg'])



```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-4-ce3b1320c030> in <cell line: 8>()
      6 # YOUR CODE HERE
      7 file_path = f'data/{SUBJECT}.csv'
----> 8 signal = pd.read_csv(file_path, header=0,names=['index','ecg'],index_col='index')
      9 
     10 

NameError: name 'pd' is not defined
```

<!-- grade_id: cell-54743a417aa594bc, points: 0, grade: false, solution: true -->
```python
# Load the labeled R-peaks an store them in pd.DataFrame called `peaks`
# YOUR CODE HERE
R_peaks_file = 'data/103_all.csv'
peaks = pd.read_csv(R_peaks_file, header=0,names=['index','R'],index_col='index')
```

Output:
```
/usr/local/lib/python3.10/dist-packages/pandas/io/parsers/readers.py in _make_engine(self, f, engine)
   1733                 if "b" not in mode:
   1734                     mode += "b"
-> 1735             self.handles = get_handle(
   1736                 f,
   1737                 mode,

/usr/local/lib/python3.10/dist-packages/pandas/io/common.py in get_handle(path_or_buf, mode, encoding, compression, memory_map, is_text, errors, storage_options)
    854         if ioargs.encoding and "b" not in ioargs.mode:
    855             # Encoding
--> 856             handle = open(
    857                 handle,
    858                 ioargs.mode,

FileNotFoundError: [Errno 2] No such file or directory: 'data/103_all.csv'
```

<!-- grade_id: cell-877b314122cf5c63, points: 4, grade: true, solution: false -->
```python
assert peaks.iloc[0, 0] == 265, 'The first R peak is found wrong.'
print('The first R peak is found correct.')
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-3-6ab31da6c0db> in <cell line: 1>()
----> 1 assert peaks.iloc[0, 0] == 265, 'The first R peak is found wrong.'
      2 print('The first R peak is found correct.')

NameError: name 'peaks' is not defined
```

```python
plt.figure()
signal.plot()

# r peaks
pos = peaks['R']/sampling_rate
plt.plot(pos, signal.loc[pos], 'r*')

plt.xlabel('time [s]')
plt.xlim(200, 205);
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-5-eec9378fffb5> in <cell line: 1>()
----> 1 plt.figure()
      2 signal.plot()
      3 
      4 # r peaks
      5 pos = peaks['R']/sampling_rate

NameError: name 'plt' is not defined
```

<!-- grade_id: cell-35a744c028a03b2c, points: 0, grade: false, solution: true -->
```python
def extract_qrst_complexes(signal, r_peaks, sampling_rate=360):
    """Extract all qrst complexes from a continous signal.

    Cuts out a region of -50ms to 250ms around each provided R-peak

    Args:
        signal: np.array of ecg signal
        r_peaks: list of indices indicating the postion of R-peaks in the signal
        sampling_rate: Ssampling rate of the recording

    Returns:
        qrst: np.array with fixed number of culumns and one qrst complex per row
    """
    signal = np.array(signal)
    # YOUR CODE HERE
    window_start = int(-0.05 * sampling_rate)
    window_end = int(0.25 * sampling_rate)
        # Initialize an empty array to store qrst complexes
    qrst = np.zeros((len(r_peaks), window_end - window_start))

    for i,r_peak in enumerate(r_peaks):
      start_index = max(0, r_peak + window_start)
      end_index = min(len(signal), r_peak + window_end)

      qrst[i,:] = signal[start_index:end_index,0]


    return qrst
```

```python
qrst_complexes = extract_qrst_complexes(signal, r_peaks=peaks['R'])

```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-9-14c92206ed6c> in <cell line: 1>()
----> 1 qrst_complexes = extract_qrst_complexes(signal, r_peaks=peaks['R'])
      2 print(sum(sum(qrst_complexes[0:1040][0:1400])) - 2836.8060839440473)

NameError: name 'signal' is not defined
```

<!-- grade_id: cell-ddc6759892bf83bc, points: 6, grade: true, solution: false -->
```python
assert abs(sum(sum(qrst_complexes[0:1040][0:1400])) - 2836.8060839440473) < 1e-3, 'The half of the extracted qrst complexes are wrong'
print('The half of the extracted qrst complexes are correct.')
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-10-597e31e33e0e> in <cell line: 1>()
----> 1 assert abs(sum(sum(qrst_complexes[0:1040][0:1400])) - 2836.8060839440473) < 1e-3, 'The half of the extracted qrst complexes are wrong'
      2 print('The half of the extracted qrst complexes are correct.')

NameError: name 'qrst_complexes' is not defined
```

<!-- grade_id: cell-73abbbb40477df25, points: 0, grade: false, solution: true -->
```python
def find_q_peak(qrst, sampling_rate=360):
    """Find the Q peak in an array of qrst complexes extracted before.

    The Q peak is the first valley to the left of the R-Peak

    Returns:
        The relative distance from the R-Peak in samples
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    return index_q

def find_s_peak(qrst, sampling_rate=360):
    """Find the S peak in an array of qrst complexes extracted before.

    The S peak is the first valley to the right of the R-Peak

    Returns:
        The relative distance from the R-Peak in samples
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    return index_s
```

```python
peaks['Q'] = peaks['R'] - find_q_peak(qrst_complexes)
peaks['S'] = peaks['R'] + find_s_peak(qrst_complexes)
peaks.head(5)
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-11-126fc96b6e8b> in <cell line: 1>()
----> 1 peaks['Q'] = peaks['R'] - find_q_peak(qrst_complexes)
      2 peaks['S'] = peaks['R'] + find_s_peak(qrst_complexes)
      3 peaks.head(5)

NameError: name 'peaks' is not defined
```

<!-- grade_id: cell-f1751529e9be2a06, points: 4, grade: true, solution: false -->
```python
assert peaks['Q'][0] == 254, 'The first Q peak is found wrong.'
assert peaks['S'][0] == 273, 'The first S peak is found wrong.'
print('The first Q and S peaks are found correct.')

```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-12-8690a17d68f8> in <cell line: 1>()
----> 1 assert peaks['Q'][0] == 254, 'The first Q peak is found wrong.'
      2 assert peaks['S'][0] == 273, 'The first S peak is found wrong.'
      3 print('The first Q and S peaks are found correct.')

NameError: name 'peaks' is not defined
```

```python
plt.figure()
plt.plot(signal)

peak_pos = peaks[['Q', 'R', 'S']] / sampling_rate
plt.plot(peak_pos['R'], signal.loc[peak_pos['R']], 'r*')
plt.plot(peak_pos['S'], signal.loc[peak_pos['S']], 'go')
plt.plot(peak_pos['Q'], signal.loc[peak_pos['Q']], 'ks')

plt.xlabel('time [s]')
plt.xlim(200, 205);
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-13-5752dbca2986> in <cell line: 1>()
----> 1 plt.figure()
      2 plt.plot(signal)
      3 
      4 peak_pos = peaks[['Q', 'R', 'S']] / sampling_rate
      5 plt.plot(peak_pos['R'], signal.loc[peak_pos['R']], 'r*')

NameError: name 'plt' is not defined
```

## Ex. 11.2: ECG feature

To describe each extracted QRST complexes we can calculate a set of features on each complex.
These features can later be used in a ML classification.

a) Implement the following features and add each of them as a new column to the `peaks` array:

1. Compute the QRS widths as the difference between the S-peaks and the Q-peaks

2. Compute the amplitude of the QR difference

3. Compute the amplitude of the RS difference

4. Compute the area below the QRST curve as follows:

- Subtract the mean of each QRST complex from each complex.
- take the absolut of all values
- Compute the area of the resulting curve by summing up all the elements.

5. Compute the normalized scalar product of a QRS complex and the template of a normal QRS complex with the following formula:
$$
c=\frac{\sum_{n = 1}^{N} [x(n) - min(x)] [y (n) - min(y)]}{\sqrt{\sum_{n=1}^{N}[x(n) - min(x)]^2 \sum_{n=1}^{N}[y(n) - min(y)]^2 }}
$$
$x(n) = $ template (first QRS complex of the ECG signal),
$y(n) = $ the QRS complex of interest,
$N = $ amount of samples the heartbeat contains,

b) Plot the results of all features and visually inspect potential outliers.

<!-- grade_id: cell-b9805980f5b5bfae, points: 0, grade: false, solution: true -->
```python
# Implement the following feature and add each of them as a new column to the `peaks` array

# QRS-width
# Calculate the distance between the S- and the Q-"valley" in seconds. The column should be named as 'QRS_Width'
# YOUR CODE HERE
def calculate_qrs_width(q_peak_indices, s_peak_indices, sampling_rate=360):
    """Calculate the QRS-width in seconds.

    Args:
        q_peak_indices: Array of relative distances from the R-peak to the Q-peak in samples
        s_peak_indices: Array of relative distances from the R-peak to the S-peak in samples
        sampling_rate: Sampling rate of the recording

    Returns:
        qrs_width: Array of QRS-width values in seconds
    """
    qrs_width = (s_peak_indices + q_peak_indices) / sampling_rate
    return qrs_width

# QR Amplitude difference
# Calculate the amplitude difference between the R-peak and the Q-valey. The column should be named as 'QR_amplitude'
# Find Q-peak and S-peak indices
sampling_rate=360
index_q = find_q_peak(qrst_complexes, sampling_rate)
index_s = find_s_peak(qrst_complexes, sampling_rate)

qrs_width = calculate_qrs_width(index_q, index_s, sampling_rate)

peaks['QRS_Width'] = qrs_width
# RS Amplitude difference
# Calculate the amplitude difference between the R-peak and the S-valey. The column should be named as 'RS_amplitude'
### BEGIN SOLTUION
peaks['RS_amplitude'] = signal.values[peaks['R'].values] - signal.values[peaks['S'].values]

# QRST Area
# Calculate the area under each QRST complex after substracting the mean value of each complex to
# normalize the baseline. The column should be named as 'QRST_area'
# YOUR CODE HERE
mean_subtracted_qrst = qrst_complexes - np.mean(qrst_complexes, axis=1, keepdims=True)
# Take the absolute value of all values
abs_mean_subtracted_qrst = np.abs(mean_subtracted_qrst)

# Compute the area of the resulting curve by summing up all the elements
area_under_qrst = np.sum(abs_mean_subtracted_qrst, axis=1)

peaks['QRST_area'] = area_under_qrst

# Template Match
# We assume the first peak in the series is a normal QRST complex
# Substract the baseline (minimum) from every qrst complex and then calculate the normalised dot-product (a*b/(|a||b|))
# YOUR CODE HERE
def normalize_scalar_product(template, qrs_complex):
    # Step 1: Calculate minimum values
    min_x = np.min(template)
    min_y = np.min(qrs_complex)

    # Step 2: Calculate numerator
    numerator = np.sum((template - min_x) * (qrs_complex - min_y))

    # Step 3: Calculate denominators
    denom_x = np.sqrt(np.sum((template - min_x)**2))
    denom_y = np.sqrt(np.sum((qrs_complex - min_y)**2))

    # Step 4: Compute normalized scalar product
    c = numerator / (denom_x * denom_y)

    return c
template = qrst_complexes[0]
dot_product =normalize_scalar_product(template, qrst_complexes)

```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-15-034b3956c6c6> in <cell line: 24>()
     22 # Find Q-peak and S-peak indices
     23 sampling_rate=360
---> 24 index_q = find_q_peak(qrst_complexes, sampling_rate)
     25 index_s = find_s_peak(qrst_complexes, sampling_rate)
     26 

NameError: name 'qrst_complexes' is not defined
```

<!-- grade_id: cell-fd777a3b86c4a9c3, points: 6, grade: true, solution: false -->
```python
assert abs(sum(peaks['QRS_Width'][0:1040]) - 53.91388888888858) < 1e-3, 'Half of the samples for the QRS_Width are wrong'
assert abs(sum(peaks['QR_amplitude'][0:1040]) - 2496.889959835294) < 1e-3, 'Half of the samples for the QR_amplitude are wrong'
assert abs(sum(peaks['RS_amplitude'][0:1040]) - 2483.7101152334803) < 1e-3, 'Half of the samples for the RS_amplitude are wrong'
assert abs(sum(peaks['QRST_area'][0:1040]) - 32009.79782801534) < 1e-3, 'Half of the samples for the QRST_area are wrong'

```

fig, axs = plt.subplots(nrows=2, ncols=3, sharex=True, figsize=(9, 5))

for ax, feature in zip(axs.flatten(), ['QRS_Width', 'QR_amplitude', 'RS_amplitude', 'QRST_area', 'Template']):
    ax.plot(peaks['R']/sampling_rate, peaks[feature])
    ax.set_title(feature)

axs[1, 1].set_xlabel('time [s]')
plt.tight_layout()
axs[1, -1].remove()

## Ex. 11.3: Critical analysis of the pipeline

a) Look through the results and the plotted ECG signal. What are potential issues that could have effected our feature calculation?

b) How could we solve these issues?

```python

```

```python

```